# Laboratory 1 - Introduction to Programming Languages Class


* This is an introductory exercise for you to learn about using google colab and github.

* Each question has a value of 1 point

* For this laboratory each exercise has its answer, please copy the answer after the comment `# YOUR CODE HERE` in their respective question


**Question 1:** Write your basic information
    
    * Full Name:  
    * Student ID: 
    * GitHub ID: 
    * Faculty/ Major: 
    * Interest research areas: 
    * Date: 
    

<!-- grade_id: q1_ans, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE
* Full Name:  김재훈 (gimjaehun)
* Student ID: 20190
* GitHub ID: esquisse01
* Faculty/ Major: industry management.
* Interest research areas:  environment
* Date: 2021/09/08

**Question 2:** Create the variables *a* = 11 and *b* = 12. Print the sum of both variables (*a* + *b*)

*Answer:* > 
- `a = 11`
- `b = 12`
- `print (a + b)`


<!-- grade_id: q2_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
a = 11
b = 12
print (a + b)
```

Output:
```
23
```

<!-- grade_id: q2_test, points: 10, grade: true, solution: false -->
```python
"""Check that a + b sum 23"""
assert (a + b == 23)
```

**Question 3:** Create the variables *h* = 11 and *i* = 12. After, add 1 to the variable *h* and print the sum of *h* and *i*.  

*Answer:* > 
- `h = 11`
- `i = 12`
- `h = h + 1` 
- `print (h + i)`

<!-- grade_id: q3_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
h = 11
i = 12
h = h + 1
print (h + i)
```

Output:
```
24
```

<!-- grade_id: q3_test, points: 10, grade: true, solution: false -->
```python

```

**Question 4:** Create the variables *m* = 12 and *n* = 12. Next, add 5 to the variable *n* and print the rest of *m* and *n*. 

*Answer:* > 
- `m = 12`
- `n = 12`
- `n = n + 5` 
- `print (m - n)`


<!-- grade_id: q4_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
m = 12
n = 12
n = n + 5
print (m - n)
```

Output:
```
-5
```

<!-- grade_id: q4_test, points: 10, grade: true, solution: false -->
```python

```

**Question 5:** Create the variables *k* = 12 and *l* = 17. Next, substract 5 to the variable *k* and substract 10 to variable *l*; print the sum of *k* and *l*.  

*Answer:* > 
- `k = 12`
- `l = 17`
- `k = k - 5` 
- `l = l - 10` 
- `print (k + l)`

<!-- grade_id: q5_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
k = 12
l = 17
k = k - 5
l = l - 10
print (k + l)
```

Output:
```
14
```

<!-- grade_id: q5_test, points: 10, grade: true, solution: false -->
```python

```

**Question 6:** Write a program to display the current date and time. 

*Answer:* > 

- `import datetime`
- `now = datetime.datetime.now()`
- `print ("Current date and time : ")`
- `print (now.strftime("%Y-%m-%d %H:%M:%S"))`

<!-- grade_id: q6_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import datetime
now = datetime.datetime.now()
print ("Current date and time : ")
print (now.strftime("%Y-%m-%d %H:%M:%S"))
```

Output:
```
Current date and time : 
2021-09-08 07:16:42
```

<!-- grade_id: q6_test, points: 10, grade: true, solution: false -->
```python

```

**Question 7:** Write a function which accepts the radius of a circle from the user and compute the area.  

*Answer:* > 

- `from math import pi`
- `def radius (r):`
    - `   area = pi * r**2`
    - `   return area`

<!-- grade_id: q7_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from math import pi
def radius (r):
area = pi * r**2
return area
```

Output:
```
File "<ipython-input-10-4da95c90b1d8>", line 4
    area = pi * r**2
       ^
IndentationError: expected an indented block
```

<!-- grade_id: q7_test, points: 10, grade: true, solution: false -->
```python

```

**Question 8:** Connect to your google drive  

*Answer:* > 
- `from google.colab import drive`
- `drive.mount('/content/drive')`


<!-- grade_id: q8_ans, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Mounted at /content/drive
```

**Question 9:** Show in the directory you are currently in  

*Answer:* > 
- `!pwd`

<!-- grade_id: q9_ans, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
!pwd
```

Output:
```
/content
```

**Question 10:** Show the files in the directory */drive/My Drive/Colab Notebooks*  

*Answer:* > 
- `!ls 'drive/My Drive/Colab Notebooks'`


<!-- grade_id: q10_answ, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
!ls 'drive/My Drive/Colab Notebooks'
```

### Congratulations! 

### You complete the first laboratory. Please save your file and save it in GitHub in Lab 1. 